# Prova di esame (laboratorio)

La prova avrà una durata massima di due ore e porterà a una valutazione massima di 21 punti. Il voto finale si otterrà per somma del punteggio del test di teoria e della prova di laboratorio. 

La prova consiste in un insieme di quiz proposti attraverso un Jupyter Notebook (non sono ammessi altri strumenti software) da svolgere sui computer del laboratorio (non sono ammessi altri strumenti hardware). Durante la prova si può consultare il libro di testo, le slide usate a lezione, un manuale python (approvato dal docente) oppure le guide in linea suggerite da Jupyter Notebook. Non sono ammessi altri strumenti di consultazione. 

**Prima della prova**

Specificare la informazioni richieste nelle celle sottostanti e salvare il notebook con il proprio nome.

**Durante la prova**

Nelle celle markdown, cancellare il testo "YOUR ANSWER HERE" con la propria risposta. Nelle celle codice, rimuovere la linea di codice
```python
raise NotImplementedError()
```
e sostituirla con la propria implementazione.
Dopo ogni cella codice che deve contenere la risposta al quesito, esiste una cella codice con delle *asserzioni*, ossia delle istruzioni che verificano la correttezza della propria soluzione. (Queste asserzioni saranno usate per valutare la correttezza della risposta data e assegnare un punteggio.) E' possibile eseguire queste celle per verificare che la soluzione proposta sia corretta.


**Dopo la prova**

Salvare il file dopo aver terminato. Consegnare l'elaborato utilizzando la piattaforma ADA. 

## Dataset

### Dati Bond Yield
#### Descrizione

Medie mensili del rendimento di un'obbligazione aziendale con rating AAA di Moody's (in percentuale/anno).  

#### Formato

File CSV con 60 osservazioni di 2 variabili (più l'indice di riga):
- **time** rappresenta il mese e l'anno secondo il formato frazionario (p.e. 1990.0 rappresenta "gennaio 1990", 1990.5 rappresenta "luglio 1990", 1990.75 rappresenta "ottobre 1990", etc.)
- **value** rappresenta il rendimento percentuale (p.e. 8.99 rappresenta "8.99%")


## Domande

### Caricamento dati

Definire un dataframe di nome `dataset` e popolarla con le righe del file `BondYield.csv`. Utilizzate la prima colonna come indice di riga (utilizzare il parametro `index_col`)

<!-- grade_id: cell-e07bdfedfd0c707d, points: 0, grade: false, solution: true -->
```python
import pandas as pd

dataset = pd.read_csv('BondYield.csv' , index_col = 0)

dataset.head()
```

Output:
```
time  value
1  1990.000000   8.99
2  1990.083333   9.72
3  1990.166667   9.37
4  1990.250000   9.46
5  1990.333333   9.47
```

<!-- grade_id: cell-4d9c038b064ce1ba, points: 1, grade: true, solution: false -->
```python
assert dataset.shape[0] == 60   # 60 osservazioni
```

<!-- grade_id: cell-90386c6f098c4a16, points: 1, grade: true, solution: false -->
```python
assert dataset.shape[1] >= 2   # almeno 2 colonne
```

<!-- grade_id: cell-6c2343029509c845, points: 1, grade: true, solution: false -->
```python
assert dataset.shape[1] == 2   # esattamente 2 colonne se la prima è usata come indice di riga
```

### Funzioni

Definire una funzione di nome `date_fromFractional` che riceve un parametro di nome `frac_date` (assunto float) e restituisce una tupla `(month,year)` che corrisponde al mese e all'anno rappresentato da `frac_date`. Per esempio:
- `date_fromFractional(1990.5) == (7, 1990)`
- `date_fromFractional(1990.75) == (10, 1990)`

**Nota** la funzione deve restituire risultati corretti per le date presenti nel dataset

<!-- grade_id: cell-b9edb1b9636cfaec, points: 0, grade: false, solution: true -->
```python
import datetime as dt
from math import ceil

def date_fromFractional(frac_date):
 
    date = dt.date(int(frac_date) , 1 , 1) + dt.timedelta(days = ceil(365 * (frac_date % 1)))
    
    return (date.month , date.year)


print(date_fromFractional(1990.75))
```

Output:
```
(10, 1990)
```

<!-- grade_id: cell-bcc8898b993e1c6f, points: 1, grade: true, solution: false -->
```python
assert date_fromFractional(1989.0) == (1, 1989)   # mese di gennaio corretto
```

<!-- grade_id: cell-340d8e6a13361f05, points: 1, grade: true, solution: false -->
```python
assert date_fromFractional(1990.5) == (7, 1990)   # mese di luglio corretto
```

<!-- grade_id: cell-d8843b87c9326ff8, points: 1, grade: true, solution: false -->
```python
assert date_fromFractional(dataset.time.iat[1]) == (2, 1990) # febbraio 1990 come nel dataset
```

<!-- grade_id: cell-0d96f07cbe0e1681, points: 1, grade: true, solution: false -->
```python
assert all(date_fromFractional(1990+i/12) == (i+1, 1990) for i in range(12)) # tutti i mesi alla precisione di macchina
```

### Manipolazione dataframe

Sostituire la colonna `time` di `dataset` con due colonne, denominate `month` e `year` che conterranno, rispettivamente, mese e anno restituite dalla funzione `date_fromFractional` applicata al valore di `time`.

**Attenzione** Gli indici del dataset iniziano da 1, non da 0. Si suggerisce di costruire prima due Series contenenti il mese e l'anno, con gli stessi indici del dataset (l'indice del dataset si ottiene con l'attributo `dataset.index`) e poi unire le due Series al dataset. Infine si può rimuovere la colonna `time`.

*Nota:* Se la funzione `date_fromFractional` non è stata correttamente implementata, inserire a mano i valori di mese e anno.

<!-- grade_id: cell-05ecfb9a6ff1944b, points: 0, grade: false, solution: true -->
```python
time = 'time'
dataset[['month', 'year']] = [(date_fromFractional(time)[0] , date_fromFractional(time)[1]) for time in dataset.pop(time)]

dataset.head()
```

Output:
```
value  month  year
1   8.99      1  1990
2   9.72      2  1990
3   9.37      3  1990
4   9.46      4  1990
5   9.47      5  1990
```

<!-- grade_id: cell-617c143fe2e266cc, points: 1, grade: true, solution: false -->
```python
assert not dataset.month.empty   # è presente la colonna month
```

<!-- grade_id: cell-333bfd72c6b174bd, points: 1, grade: true, solution: false -->
```python
assert not dataset.year.empty    # è presente la colonna year
```

<!-- grade_id: cell-72cfde5c3ae3f858, points: 1, grade: true, solution: false -->
```python
assert (dataset.month.iat[0], dataset.year.iat[0]) == (1, 1990) # le nuove colonne hanno valori corretti
```

<!-- grade_id: cell-2c1c3be42d6593d8, points: 1, grade: true, solution: false -->
```python
assert time not in dataset.columns # la colonna `time` è stata rimossa
```

<!-- grade_id: cell-162cbd96746bc5f8, points: 1, grade: true, solution: false -->
```python
assert 0 not in dataset.index  # l'indice 0 non è presente nel dataset.
```

### Dizionari

Definire un dizionario denominato `yearly_stats` che associa a ciascun anno presente nel dataset un altro dizionario con le voci `'min'` e `'max'` che contengano rispettivamente, il minimo e il massimo valore della colonna `value` del dataset, limitatamente all'anno della voce principale. Per esempio, `yearly_stats` conterrà la voce `1990` il cui valore è un dizionario con la voce `'min'` con valore `8.99` e la voce `'max'` con valore `9.72`, corrispondenti ai rendimenti minimi e massimi per l'anno 1990 registrati nel dataset.

**Nota** Assicurarsi che gli anni siano memorizzati come voci del dizionario siano di tipo `int`.

<!-- grade_id: cell-11e52afb7be70eed, points: 0, grade: false, solution: true -->
```python
yearly_stats = dict(
    (elem , {'min' : dataset[dataset.year == elem].value.min() , 'max' : dataset[dataset.year == elem].value.max()}) 
    for elem in dataset.year
)
    
yearly_stats
```

Output:
```
{1990: {'min': 8.99, 'max': 9.72},
 1991: {'min': 8.31, 'max': 9.04},
 1992: {'min': 7.92, 'max': 8.35},
 1993: {'min': 6.66, 'max': 7.91},
 1994: {'min': 6.92, 'max': 8.68}}
```

<!-- grade_id: cell-919ab1108b7e5c9f, points: 1, grade: true, solution: false -->
```python
assert all(y in yearly_stats for y in range(1990, 1994+1))  # tutti gli anni nel dizionario
```

<!-- grade_id: cell-b695545ce260e37a, points: 1, grade: true, solution: false -->
```python
assert all(all(k in yearly_stats[y] for k in ['min','max']) for y in yearly_stats) # struttura corretta del dizionario
```

<!-- grade_id: cell-8e528a0d62435dba, points: 1, grade: true, solution: false -->
```python
assert all(yearly_stats[y]['min'] <= yearly_stats[y]['max'] for y in yearly_stats) # contenuto valido
```

<!-- grade_id: cell-96e4784024009d29, points: 0, grade: true, solution: false -->
```python
assert all(type(y) == int for y in yearly_stats)   # (Non valutato) Attenzione al tipo della chiave del dizionario
```

### JSON

Salvare il dizionario `yearly_stats` in un file denominato `yearly_stats.json` nel formato JSON.

<!-- grade_id: cell-02d301a692f012e2, points: 0, grade: false, solution: true -->
```python
import json

with open("yearly_stats.json", "w") as yearly_st:
    json.dump(yearly_stats , yearly_st)
```

<!-- grade_id: cell-eab4d76a8be52be3, points: 1, grade: true, solution: false -->
```python
with open('yearly_stats.json', 'r') as jsonfile:
    test_dict = json.load(jsonfile)
    
```

### Comprensioni di lista

Estrarre la colonna `value` del dataset e memorizzarla in una lista denominata `yields`. Da questa lista definire un'altra lista denominata `yields_diff` il cui valore in posizione `i` è `'+'` se la differenza tra i valori in posizione `i+1` e `i` in `yields` è positiva, altrimenti sarà `'-'`. Utilizzare le comprensioni di lista per calcolare il risultato. Prestare attenzione alla lunghezza della lista risultante.

<!-- grade_id: cell-69579d0eef1c9b10, points: 0, grade: false, solution: true -->
```python
yields = list(dataset.value)
yields_diff = ['+' if yields[i] - yields[i - 1] > 0 else '-' for i in range(1 , len(yields))]

yields_diff[0:10]
```

Output:
```
['+', '-', '+', '+', '-', '-', '+', '+', '-', '-']
```

<!-- grade_id: cell-78929e148a857757, points: 1, grade: true, solution: false -->
```python
assert yields_diff[0:10] == ['+', '-', '+', '+', '-', '-', '+', '+', '-', '-'] # Contenuto corretto
```

<!-- grade_id: cell-2f885b7a8663cd33, points: 1, grade: true, solution: false -->
```python
assert len(yields_diff) == len(yields)-1 # Lunghezza corretta
```

### Espressioni regolari

Convertire `yields_diff` in una stringa e memorizzare il risultato in una variabile denominata `yields_diff_str`.

<!-- grade_id: cell-ccdef60a20b285ad, points: 0, grade: false, solution: true -->
```python
yields_diff_str = ''.join(elem for elem in yields_diff)

yields_diff_str
```

Output:
```
'+-++--++-----+--+-------++------++----------++--++++-+-+++-'
```

<!-- grade_id: cell-795f0af4ea8aaa3e, points: 1, grade: true, solution: false -->
```python
assert all(c in yields_diff_str for c in '+-') and type(yields_diff_str) == str # yields_diff_str è una stringa corretta
```

Utilizzando le espressioni regolari, estrarre tutte le sottostringhe contenenti almeno 3 caratteri uguali. Memorizzare le sottostringhe in una lista denominata `subs`.

<!-- grade_id: cell-c560f4b6b6e14121, points: 0, grade: false, solution: true -->
```python
import re

subs = re.findall(r'[+]{3,}|[-]{3,}' , yields_diff_str)

subs
```

Output:
```
['-----', '-------', '------', '----------', '++++', '+++']
```

<!-- grade_id: cell-b7e62e7df5e42451, points: 1, grade: true, solution: false -->
```python
assert all(len(s)>=3 for s in subs)   # Tutte le sottostringhe hanno almeno tre caratteri
```

<!-- grade_id: cell-c939a588784f42c1, points: 1, grade: true, solution: false -->
```python
assert all(len(set(s)) == 1 for s in subs) # Ogni sottostringa ha lo stesso carattere replicato più volte.
```