# Prova di esame (laboratorio)

La prova avrà una durata massima di due ore e porterà a una valutazione massima di 21 punti. Il voto finale si otterrà per somma del punteggio del test di teoria e della prova di laboratorio. 

La prova consiste in un insieme di quiz proposti attraverso un Jupyter Notebook (non sono ammessi altri strumenti software) da svolgere sui computer del laboratorio (non sono ammessi altri strumenti hardware). Durante la prova si può consultare il libro di testo, le slide usate a lezione, un manuale python (approvato dal docente) oppure le guide in linea suggerite da Jupyter Notebook. Non sono ammessi altri strumenti di consultazione. 

**Prima della prova**

Specificare la informazioni richieste nelle celle sottostanti e salvare il notebook con il proprio nome.

**Durante la prova**

Nelle celle markdown, cancellare il testo "YOUR ANSWER HERE" con la propria risposta. Nelle celle codice, rimuovere la linea di codice
```python
raise NotImplementedError()
```
e sostituirla con la propria implementazione.
Dopo ogni cella codice che deve contenere la risposta al quesito, esiste una cella codice con delle *asserzioni*, ossia delle istruzioni che verificano la correttezza della propria soluzione. (Queste asserzioni saranno usate per valutare la correttezza della risposta data e assegnare un punteggio.) E' possibile eseguire queste celle per verificare che la soluzione proposta sia corretta.


**Dopo la prova**

Salvare il file dopo aver terminato. Consegnare l'elaborato utilizzando la piattaforma ADA. 

## Dataset

### Salari bancari
#### Descrizione

Informazioni su impiegati di una banca statunitense 

#### Formato

File CSV con 474 osservazioni su 4 variabili (più l'indice di riga):

- job: Categoria = "custodial", "admin", "manage".
- education = Scolarità (in anni)
- gender = genere
- minority = l'impiegato appartiene a una minorità?


## Domande

### Caricamento dati

Definire una lista di nome `dataset` e popolarla con le righe del file `BankWages.csv`. Attenzione: ogni elemento di `dataset` deve essere una stringa e corrispondere con una riga del file. Eliminare il primo elemento dalla lista.

<!-- grade_id: cell-fb0bcfc66e6e01a6, points: 0, grade: false, solution: true -->
```python
with open('BankWages.csv') as bank:
    dataset = list(bank.readlines())[1 :]

dataset[0:10]
```

Output:
```
['"1","manage",15,"male","no"\n',
 '"2","admin",16,"male","no"\n',
 '"3","admin",12,"female","no"\n',
 '"4","admin",8,"female","no"\n',
 '"5","admin",15,"male","no"\n',
 '"6","admin",15,"male","no"\n',
 '"7","admin",15,"male","no"\n',
 '"8","admin",12,"female","no"\n',
 '"9","admin",15,"female","no"\n',
 '"10","admin",12,"female","no"\n']
```

<!-- grade_id: cell-0fd6199488f46dbc, points: 1, grade: true, solution: false -->
```python
assert len(dataset) >= 474, "Hai caricato correttamente il dataset?"
```

<!-- grade_id: cell-15526262dfb14d63, points: 1, grade: true, solution: false -->
```python
assert len(dataset) == 474, "Hai eliminato la prima riga?"
```

<!-- grade_id: cell-53e389ec62ca545f, points: 1, grade: true, solution: false -->
```python
assert type(dataset[0]) == str, "Hai caricato correttamente gli elementi del dataset?"
```

### Manipolazione di stringhe

Modificare ciascun elemento di `dataset` per eliminare i caratteri `"` e `\n` dalla stringa. Per esempio, la stringa `'"1","manage",15,"male","no"\n'` deve diventare `'1,manage,15,male,no'`.

<!-- grade_id: cell-016aedaeb0d4cb30, points: 0, grade: false, solution: true -->
```python
dataset = [elem.replace('"' , '').replace('\n' , '') for elem in dataset]

dataset[0:10]
```

Output:
```
['1,manage,15,male,no',
 '2,admin,16,male,no',
 '3,admin,12,female,no',
 '4,admin,8,female,no',
 '5,admin,15,male,no',
 '6,admin,15,male,no',
 '7,admin,15,male,no',
 '8,admin,12,female,no',
 '9,admin,15,female,no',
 '10,admin,12,female,no']
```

<!-- grade_id: cell-173698c1a9b3a63f, points: 1, grade: true, solution: false -->
```python
assert all('"' not in s for s in dataset), "Hai eliminato tutti i doppi apici?"
```

<!-- grade_id: cell-65ea9d2c3a0cba33, points: 1, grade: true, solution: false -->
```python
assert all('\n' not in s for s in dataset), "Hai eliminato tutte le newline?"
```

<!-- grade_id: cell-8701d64e78a1bff2, points: 1, grade: true, solution: false -->
```python
assert len(dataset) >= 474, "Hai modificato correttamente il dataset?"
```

### Funzioni

Definire una funzione con nome `abbrv` con parametro `s` (si assume una stringa) che opera in questo modo:
> Se la stringa `s` non contiene caratteri numerici, ha una lunghezza maggiore a 4 caratteri e non finisce con il punto `.`, restituisce una stringa costituita dai primi tre caratteri di `s` seguiti dal punto `.`, altrimenti restituisce la stringa `s`

<!-- grade_id: cell-246f3549ea2d6410, points: 0, grade: false, solution: true -->
```python
def abbrv(s):
    
    return s[0 : 3] + '.' if len(s) > 4 and not s.endswith('.') and not True in [s[i].isdigit() for i in range(len(s))] else s


abbrv('abbreviato')
```

Output:
```
'abb.'
```

<!-- grade_id: cell-d82fddac25e2cc0c, points: 1, grade: true, solution: false -->
```python
assert abbrv('extra10')=='extra10', "Hai verificato la condizione sui caratteri?"
```

<!-- grade_id: cell-ad285abda81682d0, points: 1, grade: true, solution: false -->
```python
assert abbrv('male')=='male', "Hai verificato la condizione sulla lunghezza?"
```

<!-- grade_id: cell-bf0e6bc0979d0da9, points: 1, grade: true, solution: false -->
```python
assert abbrv('manage.')=='manage.', "Hai verificato la condizione sul punto?"
```

<!-- grade_id: cell-a0576cf39e65d155, points: 1, grade: true, solution: false -->
```python
assert abbrv('female')=='fem.', "Hai trasformato correttamente la stringa?"
```

### Programmazione funzionale

Attraverso la funzione `map` e i metodi `split` e `join` e utilizzando l'istruzione `lambda` definire una funzione `f` che scompone la stringa `s` nelle sottostringhe separate da virgola, applica a ciascuna sottostringa la funzione `abbrv` e ricompone la stringa unendo le sottostringhe con la virgola.

Per esempio: `f("1,manage,15,male,no")=="1,man.,15,male,no"`

<!-- grade_id: cell-4aa5f12869696723, points: 0, grade: false, solution: true -->
```python
def f(s):
  
    return ','.join(elem for elem in map(lambda x: abbrv(x), s.split(',')))
```

<!-- grade_id: cell-5b8e638086c74638, points: 1, grade: true, solution: false -->
```python
assert f("1,manage,15,male,no")=="1,man.,15,male,no", "Hai definito correttamente la funzione?"
```

Attraverso la funzione `map`, applicare la funzione `f` a tutti gli elementi di dataset. Restituire il risultato in un oggetto di nome `dataset_map`

<!-- grade_id: cell-f1fe14efaddfcf9f, points: 0, grade: false, solution: true -->
```python
dataset_map = map(f, dataset)
```

<!-- grade_id: cell-652081ba3e9d5a94, points: 1, grade: true, solution: false -->
```python
assert type(dataset_map) == map, "Hai creato correttamente la map?"
```

Utilizzando la funzione `filter`, filtrare da `dataset_map` solo gli elementi che hanno una lunghezza pari a 20. Memorizzare il risultato nella variabile `dataset_filter`

<!-- grade_id: cell-cdb49e2a2df3d356, points: 0, grade: false, solution: true -->
```python
dataset_filter = filter(lambda x: len(x) == 20, dataset_map)
```

<!-- grade_id: cell-180e1f7b1fdebcb6, points: 1, grade: true, solution: false -->
```python
assert type(dataset_filter) == filter, "Hai costruito correttamente il filtro?"
```

Trasformare `dataset_filter` in una lista denominata `filtered_dataset`.

<!-- grade_id: cell-faf152cf262df067, points: 0, grade: false, solution: true -->
```python
filtered_dataset = list(dataset_filter)

filtered_dataset[0:10]
```

Output:
```
['109,adm.,12,male,yes',
 '110,adm.,15,male,yes',
 '111,cus.,12,male,yes',
 '112,cus.,12,male,yes',
 '125,adm.,12,male,yes',
 '126,cus.,15,male,yes',
 '127,cus.,12,male,yes',
 '128,adm.,12,fem.,yes',
 '145,adm.,12,male,yes',
 '146,adm.,18,male,yes']
```

<!-- grade_id: cell-39a50bd8bc0e684f, points: 1, grade: true, solution: false -->
```python
assert all(len(s)==20 for s in filtered_dataset), "Hai applicato correttamente la condizione di filtro?"
```

### Espressioni regolari

Utilizzando le espressioni regolari, estrarre dalla lista `dataset` le stringhe che corrispondono a impiegati che hanno una scolarità inferiore a 10 anni (dunque una sola cifra nel campo "education"). Memorizzare la lista delle stringhe estratte nella variabile `low_education`. (Suggerimento: osservate che il campo scolarità è l'unico numerico racchiuso tra due virgole.)

<!-- grade_id: cell-6ef526e123e4a02e, points: 0, grade: false, solution: true -->
```python
import re 

low_education = [elem for elem in dataset if re.fullmatch(r'.*,\d,.*', elem)]

low_education[:10]
```

Output:
```
['4,admin,8,female,no',
 '12,admin,8,male,yes',
 '36,admin,8,female,no',
 '44,admin,8,male,no',
 '61,admin,8,male,yes',
 '65,admin,8,male,no',
 '84,admin,8,female,yes',
 '90,admin,8,female,no',
 '92,admin,8,female,no',
 '96,custodial,8,male,yes']
```

<!-- grade_id: cell-3cf64938652f67d6, points: 1, grade: true, solution: false -->
```python
assert 60>len(low_education)>50, "Hai filtrato correttamente gli elementi della lista?"
```

<!-- grade_id: cell-101764eaf9921f7a, points: 1, grade: true, solution: false -->
```python
assert all(',8,' in s for s in low_education), "Hai utilizzato l'espressione regolare corretta?"
```

<!-- grade_id: cell-c5d74c1eceab8e9b, points: 1, grade: true, solution: false -->
```python
assert all('male' in s for s in low_education), "Hai estratto l'intera stringa?"
```

Utilizzando le espressioni regolari, estrarre il genere (`male`/`female`) da ciascuna stringa di `low_education`. Memorizzare il risultato nella lista `low_education_gender`

<!-- grade_id: cell-f531dfbb5ae5dd5f, points: 0, grade: false, solution: true -->
```python
low_education_gender = [re.search(r'(fe)?male', elem).group() for elem in low_education]

low_education_gender[:10]
```

Output:
```
['female',
 'male',
 'female',
 'male',
 'male',
 'male',
 'female',
 'female',
 'female',
 'male']
```

<!-- grade_id: cell-61cc5b0b6bea2460, points: 1, grade: true, solution: false -->
```python
assert all(type(s) == str for s in low_education_gender), "Hai estratto correttamente le stringhe?"
```

<!-- grade_id: cell-69e8cca2f033ff43, points: 1, grade: true, solution: false -->
```python
assert len(low_education)==len(low_education_gender), "Ogni elemento di low_education produce un risultato!"
```

<!-- grade_id: cell-727ab594eea33a24, points: 1, grade: true, solution: false -->
```python
assert all(s=='male' or s=='female' for s in low_education_gender), "Hai utilizzato l'espressione regolare corretta?"
```

Calcola la percentuale di maschi ("`male`") in `low_education_gender` e memorizza la percentuale in una stringa `perc` con una cifra decimale (es. `32.2%`)

<!-- grade_id: cell-1ff9191e4e9277e6, points: 0, grade: false, solution: true -->
```python
pattern = r'male'
perc = f'{len([elem for elem in low_education_gender if re.fullmatch(pattern, elem)]) / len (low_education_gender):.1%}'

perc
```

Output:
```
'43.4%'
```

<!-- grade_id: cell-6a781142a9065e27, points: 1, grade: true, solution: false -->
```python
assert re.fullmatch(r'\d{2}\.\d%', perc), "Hai calcolato e formattato correttamente i risultati?"
```